## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
vert0 = (0.0, 0.0)
vert1 = (0.0, 2.0)
vert2 = (3.0, 0.0)

def area(vert0, vert1, vert2):
    if not len(vert0) == len(vert1) == len(vert2) == 2:
        return 'One or more of the vertices do not have the right dimensions'
        
    return abs((vert0[0]*(vert1[1]-vert2[1]) + vert1[0]*(vert2[1]-vert0[1]) + vert2[0]*(vert0[1]-vert1[1]))/2)

print(area(vert0, vert1, vert2))
```

Output:
```
3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def area(vertices):
    A = 0
    for i in range(len(vertices)-1):
        A += 0.5*abs(vertices[i][0]*vertices[i+1][1] - vertices[i+1][0]*vertices[i][1])
    
    return A

vertices1 = [(0.0, 0.0), (0.0, 2.0), (3.0, 0.0)]
print(area(vertices1)) # =3

vertices2 = [(0.0, 0.0), (0.0, 2.0), (2.0, 2.0), (2.0, 0.0)]
print(area(vertices2)) # =4

vertices3 = [(0.0, 0.0), (10.0, 0.0), (5.0, 5.0), (1.0, 5.0)]
print(area(vertices3)) # =35
```

Output:
```
3.0
4.0
35.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    # YOUR CODE HERE
    if len(x) != len(y):
        if len(x) > len(y):
            return 'Vector x is bigger than vector y'
        else: 
            return 'Vector y is bigger than vector x'
    
    length = len(x)
    sum_v = [0]*length
    for i in range(length):
        sum_v[i] = x[i] + y[i]
    
    return sum_v
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]

sum_v = [x[0]+x[1] for x in zip(a,b)]

print(sum_v)
```

Output:
```
[-2, 11.3, -20, 8]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
college_to_abbr = {"Emmanuel": "EM", "Fitzwilliam": "F", "Girton": "G", "Gonville and Caius": "CAI", "Homerton": "HO"}

abbr_to_college = {}
for college, abbr in college_to_abbr.items():
    abbr_to_college[abbr] = college

print(abbr_to_college)

abbreviations = []
for abbr in abbr_to_college:
    abbreviations.append(abbr)

abbreviations.sort()
print(abbreviations)
```

Output:
```
{'G': 'Girton', 'HO': 'Homerton', 'F': 'Fitzwilliam', 'CAI': 'Gonville and Caius', 'EM': 'Emmanuel'}
['CAI', 'EM', 'F', 'G', 'HO']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
colleges = {
    "Emmanuel": {
        "abbr": "EM",
        "year": 1584,
        "no_of_students": 750
    }, 
    "Fitzwilliam": {
        "abbr": "F",
        "year": 1869,
        "no_of_students": 688
    }, 
    "Girton": {
        "abbr": "G",
        "year": 1869,
        "no_of_students": 677
    }, 
    "Gonville and Caius": {
        "abbr": "CAI",
        "year": 1348,
        "no_of_students": 719
    }, 
    "Homerton": {
        "abbr": "HO",
        "year": 1768,
        "no_of_students": 1181
    }
}

sort_by_size = sorted([(value["no_of_students"], value["abbr"]) for (key, value) in colleges.items()])
print("The college with the greatest number of students is {} with {} students.".format(sort_by_size[-1][1], sort_by_size[-1][0]))

sort_by_year = sorted([(value["year"], value["no_of_students"], value["abbr"]) for (key, value) in colleges.items()])
print("The oldest college is {} with {} students, founded in {}.".format(sort_by_year[0][2], sort_by_year[0][1], sort_by_year[0][0]))
```

Output:
```
The college with the greatest number of students is HO with 1181 students.
The oldest college is CAI with 719 students, founded in 1348.
```